<?xml version='1.0' encoding='utf-8'?>
<quiz><question type="coderunner"><name><text>Question 1</text></name><questiontext format="html"><text>&lt;![CDATA[&lt;p&gt;Question:
Write an SQL query to find the names of students who are older than 20 and have a grade greater than 85.&lt;/p&gt;]]&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>1.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><answer>```sql
SELECT name
FROM students
WHERE age &gt; 20 AND grade &gt; 85;
```</answer><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>-- Testing with original db</text></testcode><stdin><text /></stdin><expected><text>EXPECTED OUTPUT HERE</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase></testcases></question><question type="coderunner"><name><text>Question 2</text></name><questiontext format="html"><text>&lt;![CDATA[&lt;p&gt;**Question:**

Find the names and ages of students who have a grade higher than 85.

**&lt;/p&gt;]]&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>1.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><answer>**

```sql
SELECT name, age
FROM students
WHERE grade &gt; 85;
``` 

This query selects the `name` and `age` columns from the `students` table where the `grade` is greater than 85.</answer><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>-- Testing with original db</text></testcode><stdin><text /></stdin><expected><text>EXPECTED OUTPUT HERE</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase></testcases></question><question type="coderunner"><name><text>Question 3</text></name><questiontext format="html"><text>&lt;![CDATA[&lt;p&gt;Question:
Write an SQL query to fetch the names of students who have a grade greater than 75 and are younger than 20 years old.&lt;/p&gt;]]&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>1.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><answer>```sql
SELECT name 
FROM students 
WHERE grade &gt; 75 AND age &lt; 20;
```</answer><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>-- Testing with original db</text></testcode><stdin><text /></stdin><expected><text>EXPECTED OUTPUT HERE</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase></testcases></question><question type="coderunner"><name><text>Question 4</text></name><questiontext format="html"><text>&lt;![CDATA[&lt;p&gt;Question:
Write an SQL query to find the names and ages of students who are older than 20 and have a grade greater than or equal to 75.&lt;/p&gt;]]&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>1.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><answer>```sql
SELECT name, age
FROM students
WHERE age &gt; 20 AND grade &gt;= 75;
```</answer><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>-- Testing with original db</text></testcode><stdin><text /></stdin><expected><text>EXPECTED OUTPUT HERE</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase></testcases></question></quiz>